---
title: Test
description: Test
---

## Sample exercise

```yaml
type: NormalExercise
key: 82db100c31
lang: python
xp: 100
skills: 2
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Sponsored by CEWIT and SSRC

```yaml
type: VideoExercise
key: 0c4a9c8d13
xp: 50
```

`@projector_key`
79673afa908d33612484b0e94408994d
